<!DOCTYPE html>
<!--
Copyright (c) 2014 Samsung Electronics Co., Ltd.

Licensed under the Apache License, Version 2.0 (the License);
you may not use this file except in compliance with the License.
You may obtain a copy of the License at

    http://www.apache.org/licenses/LICENSE-2.0

Unless required by applicable law or agreed to in writing, software
distributed under the License is distributed on an "AS IS" BASIS,
WITHOUT WARRANTIES OR CONDITIONS OF ANY KIND, either express or implied.
See the License for the specific language governing permissions and
limitations under the License.

Authors:
        Md. Samsuddin Alam <ss.alam@samsung.com>

-->
<html>
<head>
<meta charset='utf-8'>
<title>SpeechSynthesis_speak</title>
<meta charset="utf-8"/>
<script src="support/unitcommon.js"></script>
<script src="support/webspeech_common.js"></script>
<meta name="flags" content="">
<meta name="assert" content="Test checks if SpeechSynthesis interface works well">
</head>
<body>
<p>Test pass if you can heard the hello world voice.</p>
<button onclick="startSpeechSynthesis()">Click to listen</button>
<script type="text/javascript">
//==== TEST: SpeechSynthesis_speak
//==== LABEL Check if speak method of SpeechSynthesis returns value
//==== PRIORITY P1
//==== SPEC Tizen Web API:Speech:SpeechSynthesis:resume M
//==== SPEC_URL https://dvcs.w3.org/hg/speech-api/raw-file/tip/speechapi.html
//==== TEST_CRITERIA MR

//test(function () {
    function startSpeechSynthesis() {
        var webspeechMessage = new SpeechSynthesisUtterance();
        webspeechMessage.text = "Hello World";
        webspeechMessage.lang = 'en-US';
        webspeechMessage.rate = 1.2;
        returnValue = speechSyn.speak(webspeechMessage);
        assert_equals(returnValue, undefined, "returnValue should be undefined");
    }
//}, document.title);

</script>
</body>
</html>
