<!DOCTYPE html>
<!--
Copyright (c) 2012 Intel Corporation.

Redistribution and use in source and binary forms, with or without modification,
are permitted provided that the following conditions are met:

* Redistributions of works must retain the original copyright notice, this list
  of conditions and the following disclaimer.
* Redistributions in binary form must reproduce the original copyright notice,
  this list of conditions and the following disclaimer in the documentation
  and/or other materials provided with the distribution.
* Neither the name of Intel Corporation nor the names of its contributors
  may be used to endorse or promote products derived from this work without
  specific prior written permission.

THIS SOFTWARE IS PROVIDED BY INTEL CORPORATION "AS IS"
AND ANY EXPRESS OR IMPLIED WARRANTIES, INCLUDING, BUT NOT LIMITED TO, THE
IMPLIED WARRANTIES OF MERCHANTABILITY AND FITNESS FOR A PARTICULAR PURPOSE
ARE DISCLAIMED. IN NO EVENT SHALL INTEL CORPORATION BE LIABLE FOR ANY DIRECT,
INDIRECT, INCIDENTAL, SPECIAL, EXEMPLARY, OR CONSEQUENTIAL DAMAGES (INCLUDING,
BUT NOT LIMITED TO, PROCUREMENT OF SUBSTITUTE GOODS OR SERVICES; LOSS OF USE,
DATA, OR PROFITS; OR BUSINESS INTERRUPTION) HOWEVER CAUSED AND ON ANY THEORY
OF LIABILITY, WHETHER IN CONTRACT, STRICT LIABILITY, OR TORT (INCLUDING
NEGLIGENCE OR OTHERWISE) ARISING IN ANY WAY OUT OF THE USE OF THIS SOFTWARE,
EVEN IF ADVISED OF THE POSSIBILITY OF SUCH DAMAGE.

Authors:
        Ma, YueX <yuex.ma@intel.com>

-->

<html>
<head>
<title>CallHistoryEntry_type_TEL</title>
<link rel="author" title="Intel" href="http://www.intel.com" />
<link rel="help" href="https://developer.tizen.org/help/topic/org.tizen.web.device.apireference/tizen/callhistory.html" />
<meta name="flags" content="" />
<meta name="assert" content="Check if attribute type value is TEL" />
<script src="support/unitcommon.js"></script>
<script src="support/callcommon.js"></script>
</head>
<body>
<div id="log"></div>
<script type="text/javascript">
//==== TEST: CallHistoryEntry_type_TEL
//==== LABEL Check if attribute type value is TEL
//==== EXECUTION_TYPE manual
//==== SPEC Tizen Web API:Social:Callhistory:CallHistoryEntry:type A
//==== SPEC_URL https://developer.tizen.org/help/topic/org.tizen.web.device.apireference/tizen/callhistory.html
//==== TEST_CRITERIA AT
//==== PRE Make some existence PHONE NUMBER ADDRESSING call history

setup({timeout: 90000});

var t = async_test(document.title, {timeout: 90000}), onSuccess, onError, filter, expected = "TEL";

t.step(function () {

    filter = new tizen.AttributeFilter("type", "EXACTLY", expected);

    onSuccess = t.step_func(function (results) {
        assert_type(results, "array", "Results should be an array");
        assert_true(results.length > 0, "Call is not found in history");
        assert_equals(results[0].type, expected, "type value is " + expected);
        t.done();
    });

    onError = t.step_func(function (error) {
        assert_unreached("find() error callback invoked: name: " + error.name + ", msg: " + error.message);
    });

    tizen.callhistory.find(onSuccess, onError, filter);
});

</script>
</body>
</html>
