<!DOCTYPE html>
<!--
Copyright (c) 2012 Intel Corporation.

Redistribution and use in source and binary forms, with or without modification,
are permitted provided that the following conditions are met:

* Redistributions of works must retain the original copyright notice, this list
  of conditions and the following disclaimer.
* Redistributions in binary form must reproduce the original copyright notice,
  this list of conditions and the following disclaimer in the documentation
  and/or other materials provided with the distribution.
* Neither the name of Intel Corporation nor the names of its contributors
  may be used to endorse or promote products derived from this work without
  specific prior written permission.

THIS SOFTWARE IS PROVIDED BY INTEL CORPORATION "AS IS"
AND ANY EXPRESS OR IMPLIED WARRANTIES, INCLUDING, BUT NOT LIMITED TO, THE
IMPLIED WARRANTIES OF MERCHANTABILITY AND FITNESS FOR A PARTICULAR PURPOSE
ARE DISCLAIMED. IN NO EVENT SHALL INTEL CORPORATION BE LIABLE FOR ANY DIRECT,
INDIRECT, INCIDENTAL, SPECIAL, EXEMPLARY, OR CONSEQUENTIAL DAMAGES (INCLUDING,
BUT NOT LIMITED TO, PROCUREMENT OF SUBSTITUTE GOODS OR SERVICES; LOSS OF USE,
DATA, OR PROFITS; OR BUSINESS INTERRUPTION) HOWEVER CAUSED AND ON ANY THEORY
OF LIABILITY, WHETHER IN CONTRACT, STRICT LIABILITY, OR TORT (INCLUDING
NEGLIGENCE OR OTHERWISE) ARISING IN ANY WAY OUT OF THE USE OF THIS SOFTWARE,
EVEN IF ADVISED OF THE POSSIBILITY OF SUCH DAMAGE.

Authors:
        Liu,Xin <xinx.liu@intel.com>

-->

<html>
  <head>
    <title>Geolocation Test:Geolocation_watchPosition_example_repeated_position</title>
    <link rel="author" title="Intel" href="http://www.intel.com/" />
    <link rel="help" href="http://www.w3.org/TR/2012/PR-geolocation-API-20120510/#geolocation" />
    <meta name="assert" content="Check if example of requesting repeated position updates works" />
    <script src="../resources/testharness.js"></script>
    <script src="../resources/testharnessreport.js"></script>
  </head>
  <body>
    <h1>You can click the button to get repeated position</h1>
    <input type="button" onclick="buttonClickHandler()" value="stop continue to monitor the position"/>
    <div><span id="position"></span></div>
    <div><span id="position"></span></div>
    <div id="log"></div>
    <script>
      var t = async_test(document.title, { timeout: 6000 });
      setup({ timeout: 20000 });
      var watchId;
      var increase=0;

      try {
        watchId = navigator.geolocation.watchPosition(expectedSuccessCallback, unexpectedErrorCallback, { maximumAge: 60000 });
      } catch (err) {
        t.step(function () {
          assert_unreached("Error message: { " + err.message + " }");
        });
        t.done();
      }

      function expectedSuccessCallback (position) {
        var coordinates = position.coords;
        document.getElementById("position").innerHTML="Your current position:" + coordinates.latitude + "," + coordinates.longitude;
        t.done();
      }

      function unexpectedErrorCallback (error) {
        t.step(function () {
          assert_unreached("Error message: { " + error.message + " }");
        });
        t.done();
      }

      function buttonClickHandler(){
        navigator.geolocation.clearWatch(watchId);
      }
  </script>
</body>
</html>
