<!DOCTYPE html>
<!--
Copyright (c) 2012 Intel Corporation.

Redistribution and use in source and binary forms, with or without modification,
are permitted provided that the following conditions are met:

* Redistributions of works must retain the original copyright notice, this list
  of conditions and the following disclaimer.
* Redistributions in binary form must reproduce the original copyright notice,
  this list of conditions and the following disclaimer in the documentation
  and/or other materials provided with the distribution.
* Neither the name of Intel Corporation nor the names of its contributors
  may be used to endorse or promote products derived from this work without
  specific prior written permission.

THIS SOFTWARE IS PROVIDED BY INTEL CORPORATION "AS IS"
AND ANY EXPRESS OR IMPLIED WARRANTIES, INCLUDING, BUT NOT LIMITED TO, THE
IMPLIED WARRANTIES OF MERCHANTABILITY AND FITNESS FOR A PARTICULAR PURPOSE
ARE DISCLAIMED. IN NO EVENT SHALL INTEL CORPORATION BE LIABLE FOR ANY DIRECT,
INDIRECT, INCIDENTAL, SPECIAL, EXEMPLARY, OR CONSEQUENTIAL DAMAGES (INCLUDING,
BUT NOT LIMITED TO, PROCUREMENT OF SUBSTITUTE GOODS OR SERVICES; LOSS OF USE,
DATA, OR PROFITS; OR BUSINESS INTERRUPTION) HOWEVER CAUSED AND ON ANY THEORY
OF LIABILITY, WHETHER IN CONTRACT, STRICT LIABILITY, OR TORT (INCLUDING
NEGLIGENCE OR OTHERWISE) ARISING IN ANY WAY OUT OF THE USE OF THIS SOFTWARE,
EVEN IF ADVISED OF THE POSSIBILITY OF SUCH DAMAGE.

Authors:
        leizhan <zhanx.lei@intel.com>
        Marcin Kaminski <marcin.ka@samsung.com>

-->

<html>
<head>
<title>PowerManager_release_correct_check</title>
<meta name="flags" content="" />
<meta name="assert" content="Check if PowerManager's release method works" />
<script type="text/javascript" src="support/unitcommon.js"></script>
</head>
<body>
<div id="log"></div>
<script type="text/javascript">
//==== TEST: PowerManager_release_correct_check
//==== LABEL Check if release() method does what it should
//==== EXECUTION_TYPE manual
//==== STEP Click 'Run', do not touch the screen, wait until the screen is changed
//==== EXPECT Pass
//==== PRE Please set backlight time to 15 seconds
//==== SPEC Tizen Web API:System:Power:PowerManager:release M
//==== SPEC_URL https://developer.tizen.org/help/topic/org.tizen.web.device.apireference/tizen/power.html
//==== TEST_CRITERIA MMINA MAST
setup({timeout: 90000});
var t = async_test(document.title, {timeout: 90000});
t.step(function () {
    function onScreenStateChanged(previousState, changedState) {
        assert_not_equals(previousState, changedState, "Screen state should be changed.");

        t.done();
    }

    tizen.power.release("SCREEN");
    tizen.power.release("CPU");

    tizen.power.turnScreenOn();
    tizen.power.request("SCREEN", "SCREEN_DIM");
    tizen.power.setScreenStateChangeListener(onScreenStateChanged);
    tizen.power.release("SCREEN");
    setTimeout(function () {
        t.step(function () {
            assert_unreached("tizen.power.release failure, " +
                    "screen state change has not been detected!");
        });
    }, 20000);
});

</script>
</body>
</html>
