<!DOCTYPE html>
<!--
Copyright (c) 2012 Intel Corporation.

Redistribution and use in source and binary forms, with or without modification,
are permitted provided that the following conditions are met:

* Redistributions of works must retain the original copyright notice, this list
  of conditions and the following disclaimer.
* Redistributions in binary form must reproduce the original copyright notice,
  this list of conditions and the following disclaimer in the documentation
  and/or other materials provided with the distribution.
* Neither the name of Intel Corporation nor the names of its contributors
  may be used to endorse or promote products derived from this work without
  specific prior written permission.

THIS SOFTWARE IS PROVIDED BY INTEL CORPORATION "AS IS"
AND ANY EXPRESS OR IMPLIED WARRANTIES, INCLUDING, BUT NOT LIMITED TO, THE
IMPLIED WARRANTIES OF MERCHANTABILITY AND FITNESS FOR A PARTICULAR PURPOSE
ARE DISCLAIMED. IN NO EVENT SHALL INTEL CORPORATION BE LIABLE FOR ANY DIRECT,
INDIRECT, INCIDENTAL, SPECIAL, EXEMPLARY, OR CONSEQUENTIAL DAMAGES (INCLUDING,
BUT NOT LIMITED TO, PROCUREMENT OF SUBSTITUTE GOODS OR SERVICES; LOSS OF USE,
DATA, OR PROFITS; OR BUSINESS INTERRUPTION) HOWEVER CAUSED AND ON ANY THEORY
OF LIABILITY, WHETHER IN CONTRACT, STRICT LIABILITY, OR TORT (INCLUDING
NEGLIGENCE OR OTHERWISE) ARISING IN ANY WAY OUT OF THE USE OF THIS SOFTWARE,
EVEN IF ADVISED OF THE POSSIBILITY OF SUCH DAMAGE.

Authors:

        Zhang, GeX <gex.zhang@intel.com>
        Tomasz Paciorek <t.paciorek@samsung.com>

-->
<html>
<head>
<title>ContactGroup_id_attribute</title>
<link rel="author" title="Intel" href="http://www.intel.com/"/>
<link rel="help" href="https://developer.tizen.org/help/topic/org.tizen.web.device.apireference/tizen/contact.html" />
<meta name="assert" content="Check that the attribute id is readonly"/>
<meta name="flags" content="" />
<script type="text/javascript" src="support/unitcommon.js"></script>
<script type="text/javascript" src="support/contact_common.js"></script>
</head>
<body>
<div id="log"></div>
<script type="text/javascript">
//==== TEST: ContactGroup_id_attribute
//==== LABEL Check ContactGroup for id attribute
//==== SPEC Tizen Web API:Social:Contact:ContactGroup:id A
//==== SPEC_URL https://developer.tizen.org/help/topic/org.tizen.web.device.apireference/tizen/contact.html
//==== TEST_CRITERIA AE ARO ADV AT
test(function () {
    var contactGroup, newValue = 1, addressBook;
    addressBook = tizen.contact.getDefaultAddressBook();
    contactGroup = new tizen.ContactGroup("Family", TEST_RING_1, TEST_IMAGE_1);
    assert_own_property(contactGroup, "id", "ContactGroup doesn't own id property.");
    assert_equals(contactGroup.id, null, "id attribute.");
    contactGroup.id = newValue;
    assert_equals(contactGroup.id, null, "id attribute.");
    addressBook.addGroup(contactGroup);
    assert_type(contactGroup.id, "string", "Check the group id type.");
    addressBook.removeGroup(contactGroup.id);
}, document.title);

</script>
</body>
</html>
