<!DOCTYPE html>
<!--
Copyright (c) 2013 Intel Corporation.

Redistribution and use in source and binary forms, with or without modification,
are permitted provided that the following conditions are met:

* Redistributions of works must retain the original copyright notice, this list
  of conditions and the following disclaimer.
* Redistributions in binary form must reproduce the original copyright notice,
  this list of conditions and the following disclaimer in the documentation
  and/or other materials provided with the distribution.
* Neither the name of Intel Corporation nor the names of its contributors
  may be used to endorse or promote products derived from this work without
  specific prior written permission.

THIS SOFTWARE IS PROVIDED BY INTEL CORPORATION "AS IS"
AND ANY EXPRESS OR IMPLIED WARRANTIES, INCLUDING, BUT NOT LIMITED TO, THE
IMPLIED WARRANTIES OF MERCHANTABILITY AND FITNESS FOR A PARTICULAR PURPOSE
ARE DISCLAIMED. IN NO EVENT SHALL INTEL CORPORATION BE LIABLE FOR ANY DIRECT,
INDIRECT, INCIDENTAL, SPECIAL, EXEMPLARY, OR CONSEQUENTIAL DAMAGES (INCLUDING,
BUT NOT LIMITED TO, PROCUREMENT OF SUBSTITUTE GOODS OR SERVICES; LOSS OF USE,
DATA, OR PROFITS; OR BUSINESS INTERRUPTION) HOWEVER CAUSED AND ON ANY THEORY
OF LIABILITY, WHETHER IN CONTRACT, STRICT LIABILITY, OR TORT (INCLUDING
NEGLIGENCE OR OTHERWISE) ARISING IN ANY WAY OUT OF THE USE OF THIS SOFTWARE,
EVEN IF ADVISED OF THE POSSIBILITY OF SUCH DAMAGE.

Authors:
        Zhang, Ge <gex.zhang@intel.com>

-->

<html lang="en">
  <head>
    <title>Capability Test: caps_displayresolution</title>
    <link rel="author" title="Intel" href="http://www.intel.com"/>
    <link rel="help" href="https://developer.tizen.org/help/topic/org.tizen.web.device.apireference/tizen/systeminfo.html"/>
    <meta name="flags" content=""/>
    <meta name="assert" content="Check if the device supports a display resolution of 480x800 (WVGA) or 720x1280 (HD), or a minimum resolution of 320x480 (HVGA)"/>
    <meta charset="utf-8"/>
    <script src="support/caps.js"></script>
  </head>
  <body>
    <div id="log"></div>
    <script>
        //It is strongly recommended to use a display resolution of 480x800 (WVGA) or 720x1280 (HD) for a mobile device implementations.
        //The Tizen reference implementation has only been validated with these two display resolutions
        var t = async_test(document.title);
        setup({timeout: 90000});
        if (is_caps_supported_by_system_info("screenSize480_800") || is_caps_supported_by_system_info("screenSize720_1280")) {
            t.done();
        } else {
            try {
                tizen.systeminfo.getPropertyValue("DISPLAY", getPropertyValueSuccess, getPropertyValueError);
            } catch (err) {
                t.step(function () {
                    assert_unreached("Error message: " + err.message);
                });
                t.done();
            }
        }

        function getPropertyValueSuccess(property){
            t.step(function () {
                assert_greater_than_equal(property.resolutionWidth, 320, "Display width is less than minimum screen width.");
                assert_greater_than_equal(property.resolutionHeight, 480, "Display height is less than minimum screen height.");
            });
            t.done();
        }

        function getPropertyValueError(error){
            t.step(function () {
                assert_unreached("Error message: " + error.message);
            });
            t.done();
        }
    </script>
  </body>
</html>
