<!DOCTYPE html>
<!--
Copyright (c) 2013 Intel Corporation.

Redistribution and use in source and binary forms, with or without modification,
are permitted provided that the following conditions are met:

* Redistributions of works must retain the original copyright notice, this list
  of conditions and the following disclaimer.
* Redistributions in binary form must reproduce the original copyright notice,
  this list of conditions and the following disclaimer in the documentation
  and/or other materials provided with the distribution.
* Neither the name of Intel Corporation nor the names of its contributors
  may be used to endorse or promote products derived from this work without
  specific prior written permission.

THIS SOFTWARE IS PROVIDED BY INTEL CORPORATION "AS IS"
AND ANY EXPRESS OR IMPLIED WARRANTIES, INCLUDING, BUT NOT LIMITED TO, THE
IMPLIED WARRANTIES OF MERCHANTABILITY AND FITNESS FOR A PARTICULAR PURPOSE
ARE DISCLAIMED. IN NO EVENT SHALL INTEL CORPORATION BE LIABLE FOR ANY DIRECT,
INDIRECT, INCIDENTAL, SPECIAL, EXEMPLARY, OR CONSEQUENTIAL DAMAGES (INCLUDING,
BUT NOT LIMITED TO, PROCUREMENT OF SUBSTITUTE GOODS OR SERVICES; LOSS OF USE,
DATA, OR PROFITS; OR BUSINESS INTERRUPTION) HOWEVER CAUSED AND ON ANY THEORY
OF LIABILITY, WHETHER IN CONTRACT, STRICT LIABILITY, OR TORT (INCLUDING
NEGLIGENCE OR OTHERWISE) ARISING IN ANY WAY OUT OF THE USE OF THIS SOFTWARE,
EVEN IF ADVISED OF THE POSSIBILITY OF SUCH DAMAGE.

Authors:
        Zhang, Ge <gex.zhang@intel.com>

-->

<html lang="en">
  <head>
    <title>Capability Test: caps_memorystorage</title>
    <link rel="author" title="Intel" href="http://www.intel.com"/>
    <link rel="help" href="https://developer.tizen.org/help/topic/org.tizen.web.device.apireference/tizen/systeminfo.html"/>
    <meta name="flags" content=""/>
    <meta name="assert" content="Check if a mobile device implementation MUST have at least 1GB of storage"/>
    <meta charset="utf-8"/>
    <script src="support/caps.js"></script>
  </head>
  <body>
    <div id="log"></div>
    <script>
        //A mobile device implementation MUST have at least 1GB of storage
        var t = async_test("Check if a mobile device implementation MUST have at least 1GB of storage.");
        t.step(function(){
            var onSuccessCallback = t.step_func(function(storage) {
                // hack, because there should be interface SystemInfoStorage passed to onSuccessCallback instead of SystemInfoStorageUnit[]
                for(var i=0; i<storage.length; i++){
                    if(storage[i].type == "INTERNAL"){
                        var result = storage[i].capacity/1073741824
                        assert_true(storage[i].capacity > 1073741824,"A mobile device implementation MUST have at least 1GB of storage, but it is " + result.toFixed(3) + "G");
                    }
                }
                t.done();
            });

            var onErrorCallback = t.step_func(function(error) {
                assert_unreached("getPropertyValue occurs an error: "+ error.message);
            });

            tizen.systeminfo.getPropertyValue("STORAGE", onSuccessCallback, onErrorCallback);
        });
    </script>
  </body>
</html>
