<!DOCTYPE html>
<!--
Copyright (c) 2013 Samsung Electronics Co., Ltd.

Licensed under the Apache License, Version 2.0 (the License);
you may not use this file except in compliance with the License.
You may obtain a copy of the License at

    http://www.apache.org/licenses/LICENSE-2.0

Unless required by applicable law or agreed to in writing, software
distributed under the License is distributed on an "AS IS" BASIS,
WITHOUT WARRANTIES OR CONDITIONS OF ANY KIND, either express or implied.
See the License for the specific language governing permissions and
limitations under the License.


Authors:
    JungHyuk, Park <junghyuk.park@samsung.com>

-->

<html lang="en">
  <head>
    <title>Capability Test: unsupported_download_cellular</title>
    <link rel="help" href="https://developer.tizen.org/help/topic/org.tizen.web.apireference.relatedpage/html/feature/reqfeatures.html"/>
    <link rel="help" href="https://developer.tizen.org/help/topic/org.tizen.web.device.apireference/tizen/systeminfo.html"/>
    <link rel="help" href="https://developer.tizen.org/help/topic/org.tizen.web.device.apireference/tizen/download.html"/>
    <meta name="flags" content=""/>
    <meta name="assert" content="Check if telephony is not supported on the device, it returns 'NotSupportedError' when requested"/>
    <meta charset="utf-8"/>
    <script src="support/caps.js"></script>
  </head>
  <body>
    <div id="log"></div>
    <script>
        test(function () {
            if (is_caps_supported_by_system_info("telephony") === false) {
                assert_throws({
                    name: "NotSupportedError"
                }, function () {
                    var downloadRequest = new tizen.DownloadRequest(
                        "http://download.tizen.org/tct/2_1/webapi-tizen-download-test-image-lq.png",
                        null, null, "CELLULAR", null);

                    tizen.download.start(downloadRequest);
                });
            }
        });
    </script>
  </body>
</html>