<!DOCTYPE html>
<!--
Copyright (c) 2012 Intel Corporation.

Redistribution and use in source and binary forms, with or without modification,
are permitted provided that the following conditions are met:

* Redistributions of works must retain the original copyright notice, this list
  of conditions and the following disclaimer.
* Redistributions in binary form must reproduce the original copyright notice,
  this list of conditions and the following disclaimer in the documentation
  and/or other materials provided with the distribution.
* Neither the name of Intel Corporation nor the names of its contributors
  may be used to endorse or promote products derived from this work without
  specific prior written permission.

THIS SOFTWARE IS PROVIDED BY INTEL CORPORATION "AS IS"
AND ANY EXPRESS OR IMPLIED WARRANTIES, INCLUDING, BUT NOT LIMITED TO, THE
IMPLIED WARRANTIES OF MERCHANTABILITY AND FITNESS FOR A PARTICULAR PURPOSE
ARE DISCLAIMED. IN NO EVENT SHALL INTEL CORPORATION BE LIABLE FOR ANY DIRECT,
INDIRECT, INCIDENTAL, SPECIAL, EXEMPLARY, OR CONSEQUENTIAL DAMAGES (INCLUDING,
BUT NOT LIMITED TO, PROCUREMENT OF SUBSTITUTE GOODS OR SERVICES; LOSS OF USE,
DATA, OR PROFITS; OR BUSINESS INTERRUPTION) HOWEVER CAUSED AND ON ANY THEORY
OF LIABILITY, WHETHER IN CONTRACT, STRICT LIABILITY, OR TORT (INCLUDING
NEGLIGENCE OR OTHERWISE) ARISING IN ANY WAY OUT OF THE USE OF THIS SOFTWARE,
EVEN IF ADVISED OF THE POSSIBILITY OF SUCH DAMAGE.

Authors:
        haox.li <haox.li@intel.com>

-->

<html>
  <head>
    <title>CSS3 Transitions Test: CSS3Transitions_transition-timing-function</title>
    <link rel="author" title="Intel" href="http://www.intel.com" />
    <link rel="help" href="http://www.w3.org/TR/2012/WD-css3-transitions-20120403/#transition-timing-function_tag" />
    <meta name="flags" content="" />
    <meta name="assert" content="To check if css.transition-timing-function exists" />
    <script src="support/support.js"></script>
    <script src="../resources/testharness.js"></script>
    <script src="../resources/testharnessreport.js"></script>
    <style type="text/css">
        div.test {
            background-color: blue;
            width: 100px;
            height: 100px;
            padding: 0.2em 1em;
            margin: 6em;
        }
    </style>
  </head>
  <body>
    <div id="log"></div>
    <div class="test" id="test"> Test transition-timing-function Attribute</div>
    <script>
        var div = document.querySelector("#test");
        div.style[headProp("transitionTimingFunction")] = "cubic-bezier(0.25, 0.1, 0.25, 1)";
        var test = GetCurrentStyle("transitionTimingFunction");
        var t = async_test(document.title, {timeout:500});
        t.step(function () {
            assert_equals(test, "cubic-bezier(0.25, 0.1, 0.25, 1)", "The Transitions transition-timing-function ");
        });
        t.done();
    </script>
  </body>
</html>
